{
   "evaluation_interval": "1m",
   "rule_files": [
      "rules.yml"
   ],
   "tests": [
      {
         "alert_rule_test": [
            {
               "alertname": "HighRateOf5xx",
               "eval_time": "1m",
               "exp_alerts": [
                  {
                     "exp_annotations": {
                        "description": "0.13 5xx rate on example.com"
                     },
                     "exp_labels": {
                        "host": "example.com",
                        "severity": "page"
                     }
                  }
               ]
            }
         ],
         "input_series": [
            {
               "series": "slo_errors_per_request_by_service:ratio_rate:5m{host=\"example.com\"}",
               "values": "0.13x10"
            },
            {
               "series": "slo_errors_per_request_by_service:ratio_rate:1h{host=\"example.com\"}",
               "values": "0.51x60"
            },
            {
               "series": "datetime:is_day",
               "values": "1x10"
            }
         ],
         "interval": "1m"
      },
      {
         "alert_rule_test": [
            {
               "alertname": "HighRateOf5xx",
               "eval_time": "1m",
               "exp_alerts": [ ]
            }
         ],
         "input_series": [
            {
               "series": "slo_errors_per_request_by_service:ratio_rate:5m{host=\"example.com\"}",
               "values": "0.13x10"
            },
            {
               "series": "slo_errors_per_request_by_service:ratio_rate:1h{host=\"example.com\"}",
               "values": "0.51x60"
            }
         ],
         "interval": "1m"
      },
      {
         "alert_rule_test": [
            {
               "alertname": "HighRateOf5xx",
               "eval_time": "1m",
               "exp_alerts": [ ]
            }
         ],
         "input_series": [
            {
               "series": "slo_errors_per_request_by_service:ratio_rate:5m{host=\"example.com\"}",
               "values": "0.03x10"
            },
            {
               "series": "slo_errors_per_request_by_service:ratio_rate:1h{host=\"example.com\"}",
               "values": "0.20x60"
            }
         ],
         "interval": "1m"
      },
      {
         "alert_rule_test": [
            {
               "alertname": "HighRateOf5xx",
               "eval_time": "1m",
               "exp_alerts": [ ]
            }
         ],
         "input_series": [
            {
               "series": "slo_errors_per_request_by_service:ratio_rate:5m{host=\"example.com\"}",
               "values": "0.03x10"
            },
            {
               "series": "slo_errors_per_request_by_service:ratio_rate:1h{host=\"example.com\"}",
               "values": "0.51x60"
            }
         ],
         "interval": "1m"
      }
   ]
}
