rule_files:
  - rules.yml
evaluation_interval: 1m
tests:
 - interval: 1m
   input_series:
    - series: 'slo_errors_per_request_by_service:ratio_rate:5m{host="example.com"}'
      values: '0.13x10'
    - series: 'slo_errors_per_request_by_service:ratio_rate:1h{host="example.com"}'
      values: '0.0000000051x60'
    - series: 'datetime:is_day'
      values: '0x10'
   alert_rule_test:
    - alertname: HighRateOf5xx
      eval_time: 1m
      exp_alerts:
       - exp_labels:
           severity: page
           host: example.com
         exp_annotations:
           description: '0.13 5xx rate on example.com'